# Project Euler: Problem 52

https://projecteuler.net/problem=52

It can be seen that the number, $125874$, and its double, $251748$, contain exactly the same digits, but in a different order.

Find the smallest positive integer, $x$, such that $2x$, $3x$, $4x$, $5x$, and $6x$, contain the same digits.

First, write a function `same_digits(x,y)` that returns `True` if two integers `x` and `y` have the exact same set of digits and multiplicities and `False` if they have different digits.

<!-- grade_id: projecteuler52a, points: 0, grade: false, solution: true -->
```python
def same_digits(x, y):
    """Do the integers x and y have the same digits, regardless of order."""
    return sorted(x) == sorted(y)
```

<!-- grade_id: projecteuler52b, points: 4, grade: true, solution: false -->
```python
assert same_digits('132', '321')
assert not same_digits('123', '3')
assert not same_digits('456', '0987654321')
```

Now use the `same_digits` function to solve this Euler problem. As you work on this problem, be careful to debug and test your code on small integers before trying it on the full search.

<!-- grade_id: projecteuler52c, points: 6, grade: true, solution: true -->
```python
test = lambda x, a: same_digits(str(x), str(a * x))

x = 1

while True:
    x1 = str(x)
    for a in range(2, 7):
        if not same_digits(x1, str(x * a)):
            break;
    else:
        break
    x += 1
    
x
```

Output:
```
142857
```

```python
print(x)
for a in range(2, 7):
    print(a,"* x =",a * x, "has same digits" if test(x, a) else "doesn't have same digits")
```

Output:
```
142857
2 * x = 285714 has same digits
3 * x = 428571 has same digits
4 * x = 571428 has same digits
5 * x = 714285 has same digits
6 * x = 857142 has same digits
```